### Lab 4 Assigned, Sunday October 18th; Due Sunday October 25th, 11:59pm
### Maximum grade: 25 / 22

The purpose of this lab is to:
* Get familiar with the Set Data type
* Function Definition
* Learn to write recursive functions


You are supposed to download this notebook, embed your answers to the questions in the cells below, and submit it through JupyterHub.  

### Question 1a -  Increasing Numbers in List - First Occurence(3 points)

Write a function `numIncreasing1(L)` that takes as input a list of numbers and returns a list of the __first__ sequence within that is increasing order, and has a length of at least 2.  If no increasing sequential numbers are found, (ie. the input list is in descending order) then naturally a list of just the first value is returned. 

Increasing sequence means for a number to be valid it has to be greater than the previous.

 + `[5, 2, 1, 3, 4]` returns `[1, 3, 4]`, because only the last 3 are in increasing order 
 + `[-1, 1, 4, 2, 6]` returns `[-1, 1, 4]`
 + `[5, 1, 8, 9, 10, 100]` returns `[1, 8, 9, 10, 100]` 
 + `[4, 1, 2, 3, 4, 2, 4, 6, 8]` returns `[1, 2, 3, 4]`
 
Hint: _the `break` keyword may prove useful to exit a loop_

<!-- grade_id: cell-444a4ab27463ffd7, points: 0, grade: false, solution: true -->
```python
def numIncreasing1(L):
    # Your code here!
    return 
    
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("Are my sample test cases correct?")
print(numIncreasing1([5, 2, 1, 3, 4]) == [1,3,4])
print(numIncreasing1([2, 1, 4, 5, 6]) == [1,4,5,6])
print(numIncreasing1([1,2,3,4,4,3,4,5,6,7,7])  == [1,2,3,4])
```

<!-- grade_id: cell-f3289e62bd137bda, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-65cdc4db9d70a421, points: 1, grade: true, solution: false -->
```python

```

### Question 1b: Increasing Numbers in List - Longest Streak (3 points)

Now write a function `numIncreasing2(L)` which applies similar constraints from Question 1a except returns the __longest__ running sequence of increasing numbers `L`. 


<!-- grade_id: cell-7705867d2bca8d4f, points: 0, grade: false, solution: true -->
```python
def numIncreasing2(s):
    # your code here! 
    return x
    
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("Are my sample test cases correct?")
print("#1", numIncreasing2([5,4,3,2,1]) == [5])
print("#2", numIncreasing2([1,2,3,4,5]) == [1,2,3,4,5])
print("#3", numIncreasing2([1,2,3,4,4,3,4,5,6,7,7]) == [ 3, 4, 5, 6, 7])
```

<!-- grade_id: cell-27eedc6832cd070c, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-9752c47972a6e13e, points: 1, grade: true, solution: false -->
```python

```

### Question 1c: Increasing Numbers in List - Default Arguments (2 points)

Now that you have defined the behavior for both `numIncreasing1` and `numIncreasing2`, create a "wrapper function" `numIncreasing` which takes a list `L` of numbers and a boolean `longest` which by default is equal to `False` and will trigger whether `numIncreasing1` or `numIncreasing2` is executed.

<!-- grade_id: cell-4f1662bda2bb64cb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-860f3a5db5d3d6dd, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3a32de1026482acd, points: 1, grade: true, solution: false -->
```python

```

###  Question 2 : Playing with Sets (2 points)

Write a function `uniqueSet(t)` which when given a string (with no punctuation) will return a __set__ containing the __unique__ words in the string. Capitalizations should be ignored. For example: $$Hello\equiv hello$$

<!-- grade_id: cell-7fe8c0405b18f890, points: 0, grade: false, solution: true -->
```python
#Your Code Here


# YOUR CODE HERE
raise NotImplementedError()
    

```

```python
print("Are my sample test cases correct?")
print("#1", len(uniqueSet("Born too late to have invented the internet")) == 8)
test2 = uniqueSet("Born too early to have entered the")
test2.add("matrix")
print("#2", len(test2) == 8)
test3 = uniqueSet("Born right on time")
test3.update(["to", "browse", "Dank", "Memes"])
print("#3", len(test3) == 8)
```

<!-- grade_id: cell-0c9024dda3538da3, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0aa6c8c1f907d8f5, points: 1, grade: true, solution: false -->
```python

```

###  Question 3: Partition the Values (3 points)

Given two lists  of integers `listA` and `listB`, write a function `partitionLists` that takes two list arguments and one optional boolean parameter `largerFirst` which defaults to `False`. 

Your function should return a tuple containing two sorted lists, where the first list is the same length as list A, and the second is the length of list B.  All the elements in list A must be __less than or equal__ to those in list B, and the lists should be by default __ascending__ order.  

If `largerFirst` is `True`, then the elements in list A must be __greater than or equal to__ those in list B, and the lists should be in __descending__ order.

<!-- grade_id: cell-150c7a63ffd4c1db, points: 0, grade: false, solution: true -->
```python
#Your code here
    
    
# YOUR CODE HERE
raise NotImplementedError()

```

```python
print("Are my sample test cases correct?")
print("#1", partitionLists([8,2,5,10],[4,5,6,3,7,5,3]) ==([2, 3, 3, 4], [5, 5, 5, 6, 7, 8, 10]))
print("#2", partitionLists([1,4,56,723,5,7,34,1],[4,5,543,6,3,45,7,5,432,3,65]) ==([1, 1, 3, 3, 4, 4, 5, 5], [5, 6, 7, 7, 34, 45, 56, 65, 432, 543, 723]))
print("#3", partitionLists([1,4,56,723,5,7,34,1],[4,5,543,6,3,45,7,5,432,3,65], True) == ([723, 543, 432, 65, 56, 45, 34, 7], [7, 6, 5, 5, 5, 4, 4, 3, 3, 1, 1]))

```

<!-- grade_id: cell-4c236c2200b391d8, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-d1242aa57f1e409f, points: 1, grade: true, solution: false -->
```python

```

### Question 4: Recursive Combinations (3 points)

Given a set of `n` items, we can choose `k` of them in $_nC_k$ ways. That is, if we have `n=3` items (for instance, an apple, an orange, and a pear) and we want to select `k=2` of them, there are $_3C_2 = 3$  ways to do this (an apple and an orange, an orange and a pear, or an apple and a pear).

We can compute $_nC_k$ recursively:

$$ _nC_k = \begin{cases} 1, & \mbox{if } k = 0, \\ 0, & \mbox{if } n < k, \\ _{n-1}C_{k-1} + _{n-1}C_{k}, & \mbox{otherwise}. \end{cases} $$

Write a function choose(n,k) that recursively computes the number of ways you can choose `k` out of `n` items.

<!-- grade_id: cell-97ff437c45c12ddc, points: 0, grade: false, solution: true -->
```python
# your code here
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(choose(3, 2) == 3)
print(choose(5, 3) == 10)
print(choose(20, 10) == 184756)
```

<!-- grade_id: cell-e166d922f83ef13c, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-967520d88f1ddee1, points: 2, grade: true, solution: false -->
```python

```

### Question 5a: Verifying Palindromes Recursively (3 points)

A word is a palindrome if it reads the same forwards as it does backwards. For instance, `racecar` and `kayak` are both palindromes, whereas `computer` is not.

Write a function that determines whether or not a given string is a palindrome. For this question, you must write your own function and test cases. Your function should be called `is_palindrome` and accept a single string as input, and return the boolean value `True` or `False` as output. Your function does not need to check anything with the input, you can assume that this function will only be passed entirely lowercase strings.

Use the following recursive algorithm to check if a string is a palindrome: If the string is 0 or 1 items long, it is definitely a palindrome (by definition). If the string is 2 or more characters long, remove the first and last characters, check if they are equal. If they aren't, the string is not a palindrome, otherwise, check if the remainder of the string is a palindrome.

<!-- grade_id: cell-7fbe4917a4099385, points: 0, grade: false, solution: true -->
```python
# Write your answer in this cell
def is_palindrome(s):
    return
    
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(is_palindrome("featherbed") == False)
print(is_palindrome("boltonotlob") == True)
print(is_palindrome("racecar") == True)
```

<!-- grade_id: cell-a5becab389cd268f, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a8b37e88effcb7df, points: 1, grade: true, solution: false -->
```python

```

### Question 5b: Improving our Palindrome Function (1 point)

Let's improve the palindrome function so that it works with a sentence that has punctuation in it.
First, write a function **only_letters(s)** that removes all punctuation (including spaces) from a given string and converts it to lowercase.

For this question you may assume the string contains only letters, punctuation, and spaces (no numbers).


<!-- grade_id: cell-a1b9f5ebbcf0d70d, points: 0, grade: false, solution: true -->
```python
# Write your answer in this cell
from string import punctuation

def only_letters(s):
    
    return s

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(only_letters(",.12 34A") == "a")
```

<!-- grade_id: cell-a1caa4022bc7d9ca, points: 1, grade: true, solution: false -->
```python
assert only_letters("h1A 932 409 l,./.2139LO") == "hallo"
```

### Question 5c: Combining the Two (2 points)

Now, using the only_letters function, write a function **is_palindrome_plus(s)** that determines if a string is a palindrome or not, ignoring punctuation or letter case.

```python

```

```python

```

<!-- grade_id: cell-f65d34128c43a028, points: 0, grade: false, solution: true -->
```python
# Write your answer in this cell
def is_palindrome_plus(s): 
    return
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(is_palindrome_plus("Do geese see God?") == True)
print(is_palindrome(only_letters('go hang a salami, i\'m a lasagna hog!')) == True)
```

<!-- grade_id: cell-587883d1c909c10b, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f3ee0ab48ffadaf8, points: 1, grade: true, solution: false -->
```python

```

###  Bonus: Recursive Encryption (3 points)

Privacy has been a major talking point amongst many tech giants as of the past 2 years. Let's be creative and find a way to encrypt our messages to send to one another!

Define a function `encrypt` which takes 3 arguments:
* `msg` -  a non-empty string message
* `shift` - a non-negative integer
* `func` - a lambda function that takes one integer argument and returns an integer. For the math lovers (which means all of you) the type definition looks like this: $$Int \rightarrow Int$$

The secret sauce to this encryption method is that it will shift every character by some integer amount `shift`. For example 'h' shifted 5 characters is 'm'. When `shift` is positive value, the shifted char will move in the (right) direction based on nominal value. (Consult Ascii chart for more info: http://www.asciitable.com/). 

Therefore, we will go through every character `c` in the string and modify each by increasing the ordinal/nominal value by `shift` creating a new character `c'` (pronounced c-prime or /sē/yah/prīm/).

Here's the catch, the shift for the **next** character will always change, as dictated by the lambda function `func`. The lambda function should receive as argument the **ordinal value of c' (`ord(c')`)** from that recursive call and return you a new `shift` integer value... (Oh yea, do this recursively; you'll thank me later.) But in the iterative sense, `func` supplies you with the **next** `shift` value for the **next proceeding** character.

HINT: With a recursive solution, this can be as little as 2 lines of concise code; first defining a base case (string of length 1), and then computing the solution in terms of its own smaller subproblems (the remainder of string).

WINK: Use `ord` and `chr` built ins. 

<!-- grade_id: cell-f250828c3b7fcde5, points: 0, grade: false, solution: true -->
```python
#Your code here
def encrypt(msg, shift, func):
    return
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("Are my sample test cases correct?")
print("#1", encrypt("hello",1, lambda x: (x + 1) % 10 ) == "iktsu")
print("#2", encrypt("covfefe",13, lambda x: x % 13 if x % 2 == 0 else x % 3) == "pwxiehe")
print("#3", encrypt("precious",12, lambda x: x + 1) == "|ïŕƹȣʓ̉ͽ")

```

<!-- grade_id: cell-f5569831a2b5ae4f, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-057993f6a5772739, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-f058097f9b6a8fb8, points: 1, grade: true, solution: false -->
```python

```